# Task 2: Training and Tuning with Ray

## Part 1: Training with Ray Train and Xgboost
In this task, you will train a machine learning model using the preprocessed data. The goal is to train an Xgboost model to predict the user rating for a product. 

```python
# !pip install ray==2.50.0 --quiet
# !pip install xgboost==1.7.6 --quiet
# !pip install -U tensorboardx --quiet
```

```python
import ray
import psutil
ray.shutdown()
NINE_GIB = 9 * 1024 * 1024 * 1024
ray.init(object_store_memory=int(NINE_GIB))
import pandas as pd
import os
import json
import random
import numpy as np
from pathlib import Path
seed = 41
random.seed(seed)
np.random.seed(seed)

from ray.train.xgboost import XGBoostTrainer
from ray.train import ScalingConfig, RunConfig
```

Output:
```
2025-11-20 22:33:44,977	INFO worker.py:2004 -- Started a local Ray instance. View the dashboard at [1m[32mhttp://127.0.0.1:8265 [39m[22m
/home/wew068/.local/lib/python3.11/site-packages/ray/_private/worker.py:2052: FutureWarning: Tip: In future versions of Ray, Ray will no longer override accelerator visible devices env var if num_gpus=0 or num_gpus=None (default). To enable this behavior and turn off this error message, set RAY_ACCEL_ENV_VAR_OVERRIDE_ON_ZERO=0
  warnings.warn(
[36m(XGBoostTrainer pid=42707)[0m Started distributed worker processes: 
[36m(XGBoostTrainer pid=42707)[0m - (node_id=68b8eea9288f716a1d135c3f1856321e6a4d3ecbe0672f1ef55cd21a, ip=10.32.73.154, pid=42910) world_rank=0, local_rank=0, node_rank=0
[36m(RayTrainWorker pid=42910)[0m [22:33:54] task [xgboost.ray-rank=00000000]:eb8829f4171ce3dc05d976a501000000 got new rank 0
[36m(SplitCoordinator pid=43041)[0m Registered dataset logger for dataset train_1_0
[36m(SplitCoordinator pid=43041)[0m Starting exec
```

```python
# load the preprocessed dataset as dense vectors in the parquet format
train_data_path=os.path.expanduser("~/public/pa3/ml_features_train.parquet")
train_data = ray.data.read_parquet(train_data_path)
```

Output:
```
Parquet dataset sampling 0:   0%|          | 0.00/1.00 [00:00<?, ? file/s]
2025-11-20 22:33:48,082	INFO parquet_datasource.py:699 -- Estimated parquet encoding ratio is 4.203.
2025-11-20 22:33:48,084	INFO parquet_datasource.py:759 -- Estimated parquet reader batch size at 677868 rows
```

### Training
Train a regression model (MSE objective) to predict `overall` using `XGBoostTrainer`.
- Objective: regression with squared error (`reg:squarederror`)
- Model params: `max_depth=3`, `eta=0.3`, others default.
- Use Ray Train with a suitable `ScalingConfig`; and you can set 2 CPUs per worker.
- After training, run inference on the test set.

Note: Ray will by default try to store results in `~/ray_results`. This can throw permission errors in DataHub, so you can change the location to `~/private/ray_results` using [RunConfig](https://docs.ray.io/en/latest/train/api/doc/ray.train.RunConfig.html).

<!-- grade_id: cell-75f2ec9ff54d45a6, points: 0, grade: false, solution: true -->
```python
# create the XGBoost trainer for regression on the "overall" label using the built-in trainer
import xgboost as xgb

trainer = XGBoostTrainer(
    label_column="overall",
    params={
        "objective": "reg:squarederror",
        "max_depth": 3,
        "eta": 0.3,
    },
    scaling_config=ScalingConfig(
        num_workers=1,
        use_gpu=False,
        resources_per_worker={"CPU": 2},
    ),
    run_config=RunConfig(storage_path="~/private/ray_results"),
    datasets={"train": train_data},
)
```

```python
result = trainer.fit()
```

Output:
```
2025-11-20 22:33:48,689	INFO tune.py:616 -- [output] This uses the legacy output and progress reporter, as Jupyter notebooks are not supported by the new engine, yet. For more information, please see https://github.com/ray-project/ray/issues/36949
== Status ==
Current time: 2025-11-20 22:33:50 (running for 00:00:00.12)
Using FIFO scheduling algorithm.
Logical resource usage: 3.0/8 CPUs, 0/0 GPUs
Result logdir: /tmp/ray/session_2025-11-20_22-33-41_618501_40643/artifacts/2025-11-20_22-33-48/XGBoostTrainer_2025-11-20_22-33-48/driver_artifacts
Number of trials: 1/1 (1 PENDING)


== Status ==
Current time: 2025-11-20 22:33:55 (running for 00:00:05.18)
Using FIFO scheduling algorithm.
Logical resource usage: 3.0/8 CPUs, 0/0 GPUs
Result logdir: /tmp/ray/session_2025-11-20_22-33-41_618501_40643/artifacts/2025-11-20_22-33-48/XGBoostTrainer_2025-11-20_22-33-48/driver_artifacts
Number of trials: 1/1 (1 RUNNING)


(pid=43041) Running 0: 0.00 row [00:00, ? row/s]
(pid=43041) - ReadParquet->SplitBlo
```

```python
print(result)
```

Output:
```
Result(
  metrics={'train-rmse': 1.120708220781679},
  path='/home/wew068/private/ray_results/XGBoostTrainer_2025-11-20_22-33-48/XGBoostTrainer_fb7d8_00000_0_2025-11-20_22-33-50',
  filesystem='local',
  checkpoint=Checkpoint(filesystem=local, path=/home/wew068/private/ray_results/XGBoostTrainer_2025-11-20_22-33-48/XGBoostTrainer_fb7d8_00000_0_2025-11-20_22-33-50/checkpoint_000000)
)
```

## Analyzing test data performance

Next, use the trained model to generate predictions on test data. Calculate the root mean square error (RMSE) of
the test predictions and report it in the output. 

For this task, we will make use of [`map_batches`](https://docs.ray.io/en/latest/data/api/doc/ray.data.Dataset.map_batches.html) to make a stateful transformation of the test data. 

```python
test_data_path=os.path.expanduser("~/public/pa3/ml_features_test.parquet")
test_dataset= ray.data.read_parquet(test_data_path)
```

Output:
```
Parquet dataset sampling 0:   0%|          | 0.00/1.00 [00:00<?, ? file/s]
2025-11-20 22:34:08,733	INFO parquet_datasource.py:699 -- Estimated parquet encoding ratio is 3.786.
2025-11-20 22:34:08,734	INFO parquet_datasource.py:759 -- Estimated parquet reader batch size at 677868 rows
```

```python
model = trainer.get_model(result.checkpoint)
```

<!-- grade_id: cell-fd7e85b107f63b98, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from ray.train import Checkpoint
import xgboost
import math

class Predictor:

    def __init__(self, checkpoint: Checkpoint):
        # Load model from the provided checkpoint
        self.model = XGBoostTrainer.get_model(checkpoint)
        self.label_col = "overall"

    def __call__(self, batch: pd.DataFrame) -> pd.DataFrame:
        """
        1. Get the predictions on a batch of data for an xgboost model.
        2. Return the squared errors for each entry using the label column
        """
        batch_features, batch_labels =  batch.loc[:, batch.columns != self.label_col], batch[[self.label_col]]
        dmatrix = xgboost.DMatrix(batch_features)
        batch["predictions"] = self.model.predict(dmatrix)
        errors = (batch["predictions"] - batch_labels[self.label_col])**2
        return {"se": errors}

def predict_xgboost(test_dataset, result):
    """
    Obtains the predictions for a test dataset given a `ray.train.Result` object and returns the squared errors for each entry
    Hint: ds.map_batches()
    """
    # create predictor from the result checkpoint
    predictor = Predictor(result.checkpoint)

    # apply predictor to test dataset (batch-wise) and collect squared errors
    se_ds = test_dataset.map_batches(predictor, batch_format="pandas")
    se_df = se_ds.to_pandas()

    # return squared errors as a 1-D numpy array
    squared_errors = se_df["se"].to_numpy()
    return squared_errors
```

<!-- grade_id: cell-709a7156a1c24be1, points: 0, grade: false, solution: true -->
```python
# get the root mean squared error for the test dataset using the result.
# Save the test rmse in `test_rmse` 

# Use the predictor implemented earlier to get squared errors, then compute RMSE
squared_errors = predict_xgboost(test_dataset, result)
test_rmse = float(np.sqrt(np.mean(squared_errors)))

# write to file
res_2_1 = {"test_rmse": test_rmse, 
          "train_rmse": result.metrics["train-rmse"]}
```

Output:
```
2025-11-20 22:34:09,895	INFO logging.py:293 -- Registered dataset logger for dataset dataset_5_0
2025-11-20 22:34:09,926	INFO streaming_executor.py:159 -- Starting execution of Dataset dataset_5_0. Full logs are in /tmp/ray/session_2025-11-20_22-33-41_618501_40643/logs/ray-data
2025-11-20 22:34:09,927	INFO streaming_executor.py:160 -- Execution plan of Dataset dataset_5_0: InputDataBuffer[Input] -> TaskPoolMapOperator[ReadParquet] -> TaskPoolMapOperator[MapBatches(Predictor)]
Running 0: 0.00 row [00:00, ? row/s]
- ReadParquet->SplitBlocks(75) 1: 0.00 row [00:00, ? row/s]
- MapBatches(Predictor) 2: 0.00 row [00:00, ? row/s]
2025-11-20 22:34:13,539	INFO streaming_executor.py:279 -- ✔️  Dataset dataset_5_0 execution finished in 3.61 seconds
```

<!-- grade_id: cell-26a7991b985e5a27, points: 20, grade: true, solution: false -->
```python
expected_path = Path("~/public/pa3/expected_2_1.json").expanduser()
with open(expected_path) as expected_file:
    expected = json.load(expected_file)

assert math.isclose(float(expected["train_rmse"]), float(res_2_1["train_rmse"]), abs_tol=0.01), \
    f"train_rmse mismatch: expected {expected['train_rmse']} vs {res_2_1['train_rmse']} (±0.01)"

assert math.isclose(float(expected["test_rmse"]), float(res_2_1["test_rmse"]), abs_tol=0.01), \
    f"test_rmse mismatch: expected {expected['test_rmse']} vs {res_2_1['test_rmse']} (±0.01)"

print("✅ Task 2.1 output matches expected within ±0.01 for RMSE.")

```

Output:
```
✅ Task 2.1 output matches expected within ±0.01 for RMSE.
```

# Part 2: Tuning with Ray Tune

Based on `XGBoostTrainer` you just implemented, now We'll tune the following hyperparameters:

1. `max_depth`
2. `eta`

You can read more about each hyperparameter in the [official docs](https://xgboost.readthedocs.io/en/stable/parameter.html). Since the overall search space is large, and our compute budget is limited, we'll focus on running 4 *trials* (or 4 instances of 2-tuples of hyperparameters) with a grid search.  Here are the values:

1. `max_depth`: $[3, 5]$ 
2. `eta`: $[0.3, 0.5]$

Steps to implement, repeated from the problem statement:
1. Create a new training and validation data from the original training data - with a random split of 75/25.
2. Train Xgboost models with 4 hyperparameter trials over the given grid using Ray Tune. [Offical Example](https://docs.ray.io/en/latest/tune/examples/tune-xgboost.html)
3. Select the best model with the lowest validation RMSE. 
4. Report the test RMSE for the best model and the lowest validation RMSE.

Make sure to use the same `ScalingConfig` as before. Restrict the number of concurrent trials to 1 for memory efficiency. Store the final `tune.ResultGrid` object in `result_grid` and the best result in the variable `best_result`.

<!-- grade_id: cell-274485071ae6e581, points: 0, grade: false, solution: true -->
```python
from ray import tune

# store your answers in these
best_result = None
result_grid = None

# Split the dataset into training and validation sets (75/25 split)
train_ds, valid_ds = train_data.train_test_split(test_size=0.25, seed=seed)

# Define the XGBoostTrainer with static configurations
trainer = XGBoostTrainer(
    label_column="overall",
    params={
        "objective": "reg:squarederror",
    },
    scaling_config=ScalingConfig(
        num_workers=1,
        use_gpu=False,
        resources_per_worker={"CPU": 2},
    ),
    run_config=RunConfig(storage_path="~/private/ray_results"),
    datasets={"train": train_ds, "validation": valid_ds},
)

# Define the Tuner for hyperparameter grid search
tuner = tune.Tuner(
    trainer,
    param_space={
        "params": {
            "max_depth": tune.grid_search([3, 5]),
            "eta": tune.grid_search([0.3, 0.5]),
        }
    },
    tune_config=tune.TuneConfig(
        max_concurrent_trials=1,
        metric="validation-rmse",
        mode="min",
    ),
)

# Execute the tuning process
result_grid = tuner.fit()

# Retrieve the best result based on validation RMSE
best_result = result_grid.get_best_result(metric="validation-rmse", mode="min")
```

Output:
```
<IPython.core.display.HTML object>
(pid=45016) Running 0: 0.00 row [00:00, ? row/s]
(pid=45016) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=45017) Running 0: 0.00 row [00:00, ? row/s]
(pid=45017) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=45573) Running 0: 0.00 row [00:00, ? row/s]
(pid=45573) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=45574) Running 0: 0.00 row [00:00, ? row/s]
(pid=45574) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=46126) Running 0: 0.00 row [00:00, ? row/s]
(pid=46126) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=46127) Running 0: 0.00 row [00:00, ? row/s]
(pid=46127) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=46678) Running 0: 0.00 row [00:00, ? row/s]
(pid=46678) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
(pid=46677) Running 0: 0.00 row [00:00, ? row/s]
(pid=46677) - split(1, equal=True) 1: 0.00 row [00:00, ? row/s]
2025-11-20 22:58:08,299	INFO tune.py:1009 -- Wrote the latest
```

```python
print(best_result)
```

Output:
```
Result(
  metrics={'train-rmse': 1.108322672118614, 'validation-rmse': 1.1108811485952415},
  path='/home/wew068/private/ray_results/XGBoostTrainer_2025-11-20_22-56-55/XGBoostTrainer_35e72_00003_3_eta=0.5000,max_depth=5_2025-11-20_22-57-49',
  filesystem='local',
  checkpoint=Checkpoint(filesystem=local, path=/home/wew068/private/ray_results/XGBoostTrainer_2025-11-20_22-56-55/XGBoostTrainer_35e72_00003_3_eta=0.5000,max_depth=5_2025-11-20_22-57-49/checkpoint_000000)
)
```

Now, 
1. Get the root mean squared error for the test dataset using the best result from the hyperparameter tuning experiments.
2. Report the validation rmse values for the best model as well as the given specific configurations

<!-- grade_id: cell-f729ed2a7b3039c6, points: 0, grade: false, solution: true -->
```python
def get_task_2_2_results(result_grid: tune.ResultGrid, best_result: ray.train.Result):
    res = {
        "test_rmse": None, # test rmse for the best model
        "valid_rmse": None, # validation rmse for the best model
        "valid_depth_5_eta_0.3": None, # validation rmse for max_depth=5, eta=0.3
        "valid_depth_3_eta_0.5": None, # validation rmse for max_depth=3, eta=0.5
    }

    # compute test rmse using the supplied best_result
    # best_result is a ray.train.Result - we can use its checkpoint to predict on the test set
    test_se = predict_xgboost(test_dataset, best_result)
    res['test_rmse'] = float(np.sqrt(np.mean(test_se)))

    # validation rmse for best model
    res['valid_rmse'] = float(best_result.metrics['validation-rmse'])

    # helper: iterate over results in the result_grid and extract validation rmse for specific configs
    # ResultGrid is iterable, yielding ray.train.Result-like objects
    for r in result_grid:
        # r.config may contain the nested 'params' mapping used by XGBoostTrainer
        cfg = getattr(r, 'config', {}) or {}
        params = cfg.get('params', {})
        md = params.get('max_depth')
        eta = params.get('eta')

        valid = None
        # metrics may be stored on r.metrics
        if hasattr(r, 'metrics'):
            valid = r.metrics.get('validation-rmse')
        elif hasattr(r, 'result') and isinstance(r.result, dict):
            valid = r.result.get('validation-rmse')

        if md == 5 and eta == 0.3:
            res['valid_depth_5_eta_0.3'] = float(valid)
        if md == 3 and eta == 0.5:
            res['valid_depth_3_eta_0.5'] = float(valid)

    return res
```

```python
res_2_2 = get_task_2_2_results(result_grid, best_result)
```

Output:
```
2025-11-20 23:02:38,392	INFO logging.py:293 -- Registered dataset logger for dataset dataset_33_0
2025-11-20 23:02:38,497	INFO streaming_executor.py:159 -- Starting execution of Dataset dataset_33_0. Full logs are in /tmp/ray/session_2025-11-20_22-33-41_618501_40643/logs/ray-data
2025-11-20 23:02:38,498	INFO streaming_executor.py:160 -- Execution plan of Dataset dataset_33_0: InputDataBuffer[Input] -> TaskPoolMapOperator[ReadParquet] -> TaskPoolMapOperator[MapBatches(Predictor)]
Running 0: 0.00 row [00:00, ? row/s]
- ReadParquet->SplitBlocks(75) 1: 0.00 row [00:00, ? row/s]
- MapBatches(Predictor) 2: 0.00 row [00:00, ? row/s]
2025-11-20 23:02:40,273	INFO streaming_executor.py:279 -- ✔️  Dataset dataset_33_0 execution finished in 1.77 seconds
```

<!-- grade_id: cell-540519f28aa086b6, points: 20, grade: true, solution: false -->
```python
expected_path = Path("~/public/pa3/expected_2_2.json").expanduser()
with open(expected_path) as expected_file:
    expected = json.load(expected_file)

for key in res_2_2.keys():
    assert math.isclose(float(expected[key]), float(res_2_2[key]), abs_tol=0.01), \
        f"{key} mismatch: expected {expected[key]} vs {res_2_2[key]} (±0.01)"

print("✅ Task 2.2 output matches expected within ±0.01 for RMSE.")

```

Output:
```
✅ Task 2.2 output matches expected within ±0.01 for RMSE.
```

```python
# shutdown!
ray.shutdown()
```